# MCL-DSCI-011 Programming in Python for Data Science 

# Assignment 2: Table Manipulation and Chaining

You can't learn technical subjects without hands-on practice. The assignments are an important part of the course. To submit this assignment you will need to make sure that you save your Jupyter notebook. 

Below are the links of 2 videos that explain:

1. [How to save your Jupyter notebook]() and,       
2. [How to answer a question in a Jupyter notebook assignment]().     

### Assignment Learning Goals:

By the end of the module, students are expected to:
- Demonstrate how to rename columns of a dataframe using [`.rename()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.rename.html).
- Create new or columns in a dataframe using [`.assign()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.assign.html) notation.
- Drop columns in a dataframe using [`.drop()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.drop.html)
- Use `df[]` notation to filter rows of a dataframe.
- Calculate summary statistics on grouped objects using [`.groupby()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.groupby.html) and [`.agg()`](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.agg.html).
- Explain when chaining is appropriate.
- Demonstrate chaining over multiple lines and verbs.

This assignment covers [Module 2](https://mcl-dsci-011-programming-in-python.netlify.com/module2) of the online course. You should complete this module before attempting this assignment.

Any place you see `...`, you must fill in the function, variable, or data to complete the code. Substitute the `None` with your completed code and answers then proceed to run the cell!

Note that some of the questions in this assignment will have hidden tests. This means that no feedback will be given as to the correctness of your solution. It will be left up to you to decide if your answer is sufficiently correct.

```python
# Import libraries needed for this lab
from hashlib import sha1
import test_assignment2 as t
import pandas as pd
import altair as alt

```

## 1. Loading files with Pandas 

In the previous module we loaded data from `csv` files using `pd.read_csv()` with limited arguments. However, most commonly you will need to create additional arguments or `pd.read_excel()` to load in a dataset to a dataframe that you can work with. `pd.read_excel()` and `pd.read_csv()` have many parameters that help read data files of different structures and formats. 

 In the `data` folder there are 8 different `.csv` and `.xlsx` files named `wine_letter.csv/.txt/.xlsx` (example- if the question is 1(f), the corresponding data is named `wine_f.csv`). Look at each of these files and use `pd.read_csv()` or `pd.read_excel()` to load these data so that they resemble the following:

|   | Grape | Origin | Alcohol | pH | Color | Aroma |
| :---: | :---: | :---: | :---: | :---: | :---: | :---: |
| **0** |  Chardonnay | Australia | 14.23 | 3.51 | White | Floral |
| **1** |  Pinot Grigio | Italy | 13.20 | 3.30 | White | Fruity |
| **2** |  Pinot Blanc | France | 13.16 | 3.16 | White | Citrus |
| **3** |  Shiraz | Chile | 14.91 | 3.39 | Red | Berry |
| **4** |  Malbec | Argentina | 13.83 | 3.28 | Red | Fruity |


**Question 1(a)**  <br> {points: 1}   

Load in the file named `wine_a.csv` using the appropriate verb `pd.read_csv()` or `pd.read_excel()` and arguments and save it as an object named `df_1a`.

<!-- grade_id: cell-e60ff99336d4d01a, points: 0, grade: false, solution: true -->
```python
df_1a = None

### BEGIN SOLUTION
# no special considerations here
df_1a = pd.read_csv('data/wine_a.csv')
### END SOLUTION


```

<!-- grade_id: cell-4931d4895213ff2d, points: 1, grade: true, solution: false -->
```python
t.test_1a(df_1a)
```

Output:
```
'success'
```

**Question 1(b)**  <br> {points: 1}   

Load in the file named `wine_b.csv` using the appropriate verb `pd.read_csv()` or `pd.read_excel()` and arguments and save it as an object named `df_1b`.

<!-- grade_id: cell-97c01391ad25a043, points: 0, grade: false, solution: true -->
```python
df_1b = None

### BEGIN SOLUTION
# Need to skip first row
df_1b = pd.read_csv('data/wine_b.csv',
                    header=1)

### END SOLUTION

```

<!-- grade_id: cell-09e0d6bb4c9185dd, points: 1, grade: true, solution: false -->
```python
t.test_1b(df_1b)
```

Output:
```
'success'
```

**Question 1(c)**  <br> {points: 1}   

Load in the file named `wine_c.csv` using the appropriate verb `pd.read_csv()` or `pd.read_excel()` and arguments and save it as an object named `df_1c`.

<!-- grade_id: cell-5bcbd26df73e99b7, points: 0, grade: false, solution: true -->
```python
df_1c = None

### BEGIN SOLUTION
# Need to read only first 5 rows
df_1c = pd.read_csv('data/wine_c.csv',
                    nrows=5)

### END SOLUTION


```

<!-- grade_id: cell-76b78585dc5069f4, points: 1, grade: true, solution: false -->
```python
t.test_1c(df_1c)
```

**Question 1(d)**  <br> {points: 1}   

Load in the file named `wine_d.csv` using the appropriate verb `pd.read_csv()` or `pd.read_excel()` and arguments and save it as an object named `df_1d`.

<!-- grade_id: cell-84c10d85922e4413, points: 0, grade: false, solution: true -->
```python
df_1d = None

### BEGIN SOLUTION
# Need to specify tab-delimited
df_1d = pd.read_csv('data/wine_d.txt',
                    delimiter='\t')
### END SOLUTION


```

<!-- grade_id: cell-463aaa00c2184211, points: 1, grade: true, solution: false -->
```python
t.test_1d(df_1d)
```

**Question 1(e)**  <br> {points: 1}   

Load in the file named `wine_e.csv` using the appropriate verb `pd.read_csv()` or `pd.read_excel()` and arguments and save it as an object named `df_1e`.

<!-- grade_id: cell-246416508bc9f501, points: 0, grade: false, solution: true -->
```python
df_1e = None

### BEGIN SOLUTION
# Need to use only the first 6 columns
df_1e = pd.read_csv('data/wine_e.csv',
                    usecols=[0, 1, 2, 3, 4, 5])

### END SOLUTION

```

<!-- grade_id: cell-1d4d04b4dfe42be6, points: 1, grade: true, solution: false -->
```python
assert 'df_1e' in locals(), "Your answer should be named 'df_1e'"

# Note that this test has been hidden intentionally.
# It will provide no feedback as to the correctness of your answers.
# Thus, it is up to you to decide if your answer is sufficiently correct.
### BEGIN HIDDEN TESTS

def hidden_1(answer):
    assert not answer is None, "Your answer does not exist. Have you passed in the correct variable?"
    assert answer.shape[0] == 5 and answer.shape[1] == 6, "Your dataframe dimensions are incorrect. Are you using the read_csv or read_excel function to load the data?"
    msg = "Your dataframe contains the incorrect columns. Are you using the use_cols parameter? \
    \nExpected ['Grape', 'Origin', 'Alcohol', 'pH', 'Colour','Aroma'], but got {0}".format(
        list(answer.columns))
    assert sorted(list(answer.columns)) == sorted(
        ['Grape', 'Origin', 'Alcohol', 'pH', 'Colour', 'Aroma']), msg
    assert "Chardonnay" in list(
        answer.Grape
    ), "Some values in your dataframe are incorrect. Do you have the correct dataframe?"
    assert sum(
        list(answer.Alcohol)
    ) == 69.33, "Some values in your dataframe are incorrect. Do you have the correct dataframe?"
    return("Success")

hidden_1(df_1e)
### END HIDDEN TESTS
```

Output:
```
'Success'
```

**Question 1(f)**  <br> {points: 1}   

Load in the file named `wine_f.csv` using the appropriate verb `pd.read_csv()` or `pd.read_excel()` and arguments and save it as an object named `df_1f`.

<!-- grade_id: cell-9a0f33c3a9ff9592, points: 0, grade: false, solution: true -->
```python
df_1f = None

### BEGIN SOLUTION
# Need to use pd.read_csv() with tab delimiter
df_1f = pd.read_csv('data/wine_f.txt',
                  delimiter='\t',
                  header=1,
                  nrows=5,
                  usecols=[1, 2, 3, 4, 5, 6])

### END SOLUTION

```

<!-- grade_id: cell-321bd65f6db7764e, points: 1, grade: true, solution: false -->
```python
t.test_1f(df_1f)
```

**Question 1(g)**  <br> {points: 1}   

Load in the file named `wine_g.csv` using the appropriate verb `pd.read_csv()` or `pd.read_excel()` and arguments and save it as an object named `df_1g`.

<!-- grade_id: cell-36bc923bb358d123, points: 0, grade: false, solution: true -->
```python
df_1g = None

### BEGIN SOLUTION
# Need to use pd.read_excel()
df_1g = pd.read_excel('data/wine_g.xlsx')

### END SOLUTION

```

<!-- grade_id: cell-bdebf9baa0b94dde, points: 1, grade: true, solution: false -->
```python
t.test_1g(df_1g)
```

**Question 1(h)**  <br> {points: 1}   

Load in the file named `wine_h.csv` using the appropriate verb `pd.read_csv()` or `pd.read_excel()` and arguments and save it as an object named `df_1h`.

<!-- grade_id: cell-962b21f176ed873e, points: 0, grade: false, solution: true -->
```python
df_1h = None

### BEGIN SOLUTION

# Need to use pd.read_excel() and specify sheet name
df_1h = pd.read_excel('data/wine_h.xlsx',
                    sheet_name='WINE')
### END SOLUTION


```

<!-- grade_id: cell-b7b92ccaceee0a48, points: 1, grade: true, solution: false -->
```python
t.test_1h(df_1h)
```


**Question 1(i)**  <br> {points: 1}   

Pandas can also load files directly using a url. Let's explore this functionality using a simple dataset showing the number of people that immigrate to Canada from the US and Britain located [here](https://raw.githubusercontent.com/plotly/datasets/master/Move%20to%20Canada.csv).

Read the 'Move to Canada' .csv file using `pd.read_csv()` and store the dataframe as the variable `df`.

<!-- grade_id: cell-293da28cdee572c7, points: 0, grade: false, solution: true -->
```python
df = None

### BEGIN SOLUTION
url = 'https://raw.githubusercontent.com/plotly/datasets/master/Move%20to%20Canada.csv'
df = pd.read_csv(url)
### END SOLUTION 
```

<!-- grade_id: cell-645f7dcf97036ef7, points: 1, grade: true, solution: false -->
```python
t.test_1i(df)
```

## 2. Vancouver Street Tree dataset <a name='2.6'></a>

The City of Vancouver, nicknamed [Raincouver](https://blogs.ubc.ca/camdrohan/2012/08/14/the-truth-about-raincouver/) or [Raincity](https://en.wikipedia.org/wiki/Nicknames_of_Vancouver) surprises many non-residents by the sheer quantity of rain it receives. To locals, this is simply the expectations of living in a rainforest. The rainfall mixed with the temperate climate provides the perfect environment to maintain numerous species of trees across the city. Every year, Vancouver hosts a [cherry blossom festival](https://www.vcbf.ca/)  where people can take walks to view the blossoming cherry trees across the city,  post cherry blossom photos and attend a Cherry Jam Concert. People more interested in grandeur can explore the [giants in Stanley Park](https://www.insidevancouver.ca/2019/06/27/where-to-see-big-trees-in-vancouver/). 
In any case, The city has provided open data to all of Vancouver's street trees. The data includes the following:



-  TREE_ID
- CIVIC_NUMBER - Street number
- STD_STREET - Street name
- GENUS_NAME 
- SPECIES_NAME 
- CULTIVAR_NAME 
- COMMON_NAME 
- ASSIGNED - Indicates whether the address is made up to associate the tree with a nearby  lot (Y=Yes or N=No)
- ROOT_BARRIER - Root barrier installed (Y = Yes, N = No)
- PLANT_AREA - B = behind sidewalk, G = in tree grate, N = no sidewalk, C = cutout, a number  indicates boulevard width in feet
- ON_STREET_BLOCK
- ON_STREET
- NEIGHBOURHOOD_NAME 
- STREET_SIDE_NAME - The street side which the tree is physically located on (Even, Odd or Median  (Med))
- HEIGHT_RANGE_ID - 0-10 for every 10 feet (e.g., 0 = 0-10 ft, 1 = 10-20 ft, 2 = 20-30 ft, and10 = 100+ ft)
- DIAMETER - in inches
- CURB - Curb presence (Y = Yes, N = No)
- DATE_PLANTED 
- Geom   


**Question 2(a)**  <br> {points: 1}  

Load the `street-trees.csv` dataset into a pandas dataframe named `trees`.     
_Hint: How is your file delimited?_

<!-- grade_id: cell-7a1b3cebb9d471e8, points: 0, grade: false, solution: true -->
```python
trees = None

### BEGIN SOLUTION
trees = pd.read_csv('data/street-trees.csv',delimiter=';')
### END SOLUTION 
```

<!-- grade_id: cell-9483fdf964a8582a, points: 1, grade: true, solution: false -->
```python
t.test_2a(trees)
```

Output:
```
'success'
```

**Question 2(b)**  <br> {points: 1}  

Create a new column named `DIAMETER_CM` that contains the each tree's diameter in cm. 
Don't forget to save the dataframe as the object `trees` to maintain the changes.    
_Note that 1 inch = 2.54cm._


<!-- grade_id: cell-d76b48ee4be626ca, points: 0, grade: false, solution: true -->
```python
CM_PER_INCH = 2.54
# trees = trees.assign(...)

### BEGIN SOLUTION
trees = trees.assign(DIAMETER_CM = trees['DIAMETER'] * CM_PER_INCH)
### END SOLUTION 
```

<!-- grade_id: cell-0436a6dbfeca97e4, points: 1, grade: true, solution: false -->
```python
t.test_2b(trees)
```

Output:
```
'success'
```

**Question 2(c)**  <br> {points: 1}  
Filter so we can concentrate our attention on the cherry trees - specifically species `SERRULATA`.   
Save the resulting dataframe in an object called cherry.


<!-- grade_id: cell-a607c416f05888c4, points: 0, grade: false, solution: true -->
```python
# cherry = trees[...[...] == ...]

### BEGIN SOLUTION
cherry = trees[trees['SPECIES_NAME'] == 'SERRULATA']
### END SOLUTION 
```

<!-- grade_id: cell-432b2048ee01b7a3, points: 1, grade: true, solution: false -->
```python
t.test_2c(cherry)
```

Output:
```
'success'
```

**Question 2(d)**  <br> {points: 1}  

Group the cherry trees by their common name. Save the results in an object called `cherry_group`.

<!-- grade_id: cell-6cece0b23a120a40, points: 0, grade: false, solution: true -->
```python
cherry_group = None 

### BEGIN SOLUTION
cherry_group = cherry.groupby('COMMON_NAME')
### END SOLUTION 
```

<!-- grade_id: cell-113ba90db47fd897, points: 1, grade: true, solution: false -->
```python
t.test_2d(cherry_group)
```

Output:
```
'success'
```

**Question 2(e)**  <br> {points: 1}  

For each type of cherry tree (`COMMON_NAME`) find the average diameter in cm and the number of the cherry trees for that cherry tree type.   
Name the new dataframe as `cherry_stats`.    
*Hint: Your dataframe should have the shape (16, 2).*     

<!-- grade_id: cell-23e8676185b31137, points: 0, grade: false, solution: true -->
```python
# cherry_stats = (...
#                 .agg(...),
#                 .loc[...])

### BEGIN SOLUTION
cherry_stats = (cherry.groupby('COMMON_NAME')
                .agg(['mean', 'count'])
                .loc[:,'DIAMETER_CM'])
cherry_stats
### END SOLUTION 
```

Output:
```
mean  count
COMMON_NAME                                
ACCOLADE CHERRY            52.768500     60
AMANOGAWA JAPANESE CHERRY  21.638559     68
CHEALS WEEPING CHERRY       8.128000      5
GREAT WHITE CHERRY         49.247778      9
JAPANESE FLOWERING CHERRY  42.627791    474
KWANZAN FLOWERING CHERRY   45.734833  10540
MIKURUMA-GAESHI CHERRY     40.640000      1
MIYAKO CHERRY              42.291000      2
PINK PERFECTION CHERRY     47.896517    522
SEIBOLDI CHERRY            43.499741    170
SHIROFUGEN CHERRY          20.795648    369
SHIROTAE(MT FUJI) CHERRY   34.349267    225
SHOGETSU JAPANESE CHERRY   71.882000      1
SNOW GOOSE CHERRY           8.158788     33
UKON JAPANESE CHERRY       43.817343    813
WEEPING JAPANESE CHERRY    59.700410     61
```

<!-- grade_id: cell-e05a3d9da4ed3d5c, points: 1, grade: true, solution: false -->
```python
t.test_2e(cherry_stats)
```

Output:
```
'success'
```

**Question 2(f)**  <br> {points: 1}  

Graph the mean diameter of `cherry_stats`.  Sort the values in ascending order before you use graph your results. Make sure you add a discriptive title and label the y-axis.        
Save the plot in an object called `cherry_graph`.   

<!-- grade_id: cell-d27bbe1f8b0c7834, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
alt.data_transformers.disable_max_rows()
cherry_graph = alt.Chart(cherry).mark_bar().encode(
    x=alt.X('COMMON_NAME', sort="y"), 
    y=alt.Y('mean(DIAMETER):Q')).properties(
    title="Mean tree diameter of different cherry species")
### END SOLUTION
cherry_graph
```

Output:
```
alt.Chart(...)
```

<!-- grade_id: cell-f58346aa6368a2f3, points: 1, grade: true, solution: false -->
```python
t.test_2f(cherry_graph)
```

Output:
```
'success'
```

## 3. The Titanic dataset <a name='2.6'></a>

The sinking of the RMS Titanic is one of the most infamous shipwrecks in history.  On April 15, 1912, during her maiden voyage, the Titanic sank after colliding with an iceberg. The file *titanic.csv* contains data of 1309 passengers who were on the ship. For each passenger, the following data are recorded:
* survival - Survival (0 = No; 1 = Yes)
* class - Passenger Class (1 = 1st; 2 = 2nd; 3 = 3rd)
* name - Name
* sex - Sex
* age - Age
* sibsp - Number of Siblings/Spouses Aboard
* parch - Number of Parents/Children Aboard
* ticket - Ticket Number
* fare - Passenger Fare
* cabin - Cabin
* embarked - Port of Embarkation (C = Cherbourg; Q = Queenstown; S = Southampton)
* boat - Lifeboat (if survived)
* body - Body number (if did not survive and the body was recovered)


_**DISCLOSURE:** We understand that one of the limitations of this dataset is that it reflects binary sex categories._

In this exercise, you will perform a number of wrangling operations to manipulate and extract subsets of the data.

**Question 3(a)**  <br> {points: 1}   


Load the `titanic.csv` dataset into a pandas dataframe named `titanic_df`.     
*Hint: Remember where we store our data*

<!-- grade_id: cell-6ca937ceb1f63f35, points: 0, grade: false, solution: true -->
```python
titanic_df = None

### BEGIN SOLUTION
titanic_df = pd.read_csv('data/titanic.csv')
### END SOLUTION 
```

<!-- grade_id: cell-8b69beeeb3a63066, points: 1, grade: true, solution: false -->
```python
t.test_3a(titanic_df)
```

Output:
```
'success'
```

**Question 3(b)**  <br> {points: 1}  


The column names `sibsp` and `parch` are not very descriptive. Rename these columns to `siblings_spouses` and `parents_children` respectively.

<!-- grade_id: cell-848a01803315ae81, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
titanic_df = titanic_df.rename(columns={'sibsp': 'siblings_spouses',
                                        'parch': 'parents_children'})
### END SOLUTION 
```

<!-- grade_id: cell-a971e94ff67950ce, points: 1, grade: true, solution: false -->
```python
t.test_3b(titanic_df)
```

Output:
```
'success'
```

**Question 3(c)**  <br> {points: 1}  

Drop the columns `body` and `cabin`, using a single line of code and save it as an object named `dropped_df`.

<!-- grade_id: cell-d822e26557aef2f2, points: 0, grade: false, solution: true -->
```python
dropped_df = None

### BEGIN SOLUTION
dropped_df = titanic_df.drop(columns=['body', 'cabin'])
### END SOLUTION
```

<!-- grade_id: cell-2349fcdcc552bdd3, points: 1, grade: true, solution: false -->
```python
t.test_3c(dropped_df)
```

Output:
```
'success'
```

**Question 3(d)**  <br> {points: 1}  
Which passengers survived (`survived = 1`) the disaster? Save the filtered dataframe in an object named `survived_titanic`. 

<!-- grade_id: cell-58d4de705145af3c, points: 0, grade: false, solution: true -->
```python
survived_titanic = None

### BEGIN SOLUTION
survived_titanic = titanic_df[titanic_df['survived'] == 1]
### END SOLUTION 
```

<!-- grade_id: cell-d337093c3836b042, points: 2, grade: true, solution: false -->
```python
assert 'survived_titanic' in locals(), "Your answer should be named 'survived_titanic'"

# Note that this test has been hidden intentionally.
# It will provide no feedback as to the correctness of your answers.
# Thus, it is up to you to decide if your answer is sufficiently correct.
### BEGIN HIDDEN TESTS
def hidden_2(answer):
    assert not answer is None, "Your answer does not exist. Have you passed in the correct variable?"
    assert sha1(str(answer.shape[0] + 9).encode('utf8')).hexdigest(
    ) == '6d5db0e809f71a43d3bada01e4c1c4d4b501b435', "The number of rows are incorrect. Are you selecting only survivors?"
    assert round(
        sum(answer.fare)
    ) == 24681, "The sum of the fares is incorrect. Are you selecting only survivors?"
    assert round(
        answer.age.mean()
    ) == 29.0, "The average survivor age is incorrect. Are you selecting only survivors?"
    assert "24160" in list(
        answer.ticket
    ), "Some survivors are missing. Did you select all the survivors?"
    return ("Success")

hidden_2(survived_titanic)
### END HIDDEN TESTS
```

Output:
```
'Success'
```

**Question 3(e)**  <br> {points: 1}  
How many people survived the titanic?       
Save this in an object named `survivors`.

<!-- grade_id: cell-33a45d51a162c2aa, points: 0, grade: false, solution: true -->
```python
survivors = None

### BEGIN SOLUTION
survivors = survived_titanic.shape[0]

### END SOLUTION 
```

<!-- grade_id: cell-18b12e16e1cc9d38, points: 1, grade: true, solution: false -->
```python
t.test_3e(survivors)
```

**Question 3(f)**  <br> {points: 1}  

Which passengers were in first class and younger than 16?   
Save the filtered dataframe in an object named `first_class_pass`. 

<!-- grade_id: cell-3a0619ef6d2af5a2, points: 0, grade: false, solution: true -->
```python
first_class_pass = None

### BEGIN SOLUTION
first_class_pass = titanic_df[(titanic_df['pclass'] == 1) & (titanic_df['age'] < 16)]
### END SOLUTION 
```

<!-- grade_id: cell-4ed4057ebf7bf52f, points: 1, grade: true, solution: false -->
```python
t.test_3f(first_class_pass)
```

**Question 3(g)**  <br> {points: 1}  

Make a new dataframe containing only the columns `age` and `fare` on the full `titanic_df` dataframe and take the mean of these columns using chaining.   
Save it as an object named  `mean_age_fare`


<!-- grade_id: cell-2e89f890b11155df, points: 0, grade: false, solution: true -->
```python
mean_age_fare= None 

### BEGIN SOLUTION
mean_age_fare = titanic_df.loc[:, ['age', 'fare']].mean()
### END SOLUTION 
```

<!-- grade_id: cell-d73c0f21bdebe3df, points: 1, grade: true, solution: false -->
```python
t.test_3g(mean_age_fare)
```

**Question 3(h)**  <br> {points: 1}  

Let's make a groupby object by the column `sex`  on our full dataframe `titanic_df` and find the average value each of each column for `male` and `female` passengers.    
Save the results in an object named `titanic_sex_mean`.

<!-- grade_id: cell-2518084eff7a7e46, points: 0, grade: false, solution: true -->
```python
titanic_sex_mean = None

### BEGIN SOLUTION
titanic_sex_mean = titanic_df.groupby(by='sex').mean()
titanic_sex_mean
### END SOLUTION 
```

<!-- grade_id: cell-c59fc52c7c1f03c4, points: 1, grade: true, solution: false -->
```python
t.test_3h(titanic_sex_mean)
```

**Question 3(i)**  <br> {points: 1}  

To answer the question, assign the letter associated with the correct answer to a variable in the code cell below:   
Which statement below is accurate?

A. The average age of male passengers is lower than the total passenger mean age.    

B. The female passenger fare average is greater than the male passenger fare average.  

C. The male fare average is greater than than the total passenger fare average. 


*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between `""`, assign the correct answer to an object called `answer3_i`*

<!-- grade_id: cell-b14e8438781ba083, points: 0, grade: false, solution: true -->
```python
answer3_i = None

### BEGIN SOLUTION
answer3_i = "B"
### END SOLUTION
```

<!-- grade_id: cell-884d987b24fb9b40, points: 1, grade: true, solution: false -->
```python
t.test_3i(answer3_i)
```

**Question 3(j)**  <br> {points: 1}  

What was the highest and lowest fair paid by male and female passengers? Use `.agg()` and save the dataframe as object `high_low_fare`.  


<!-- grade_id: cell-01ad2d8d34762607, points: 0, grade: false, solution: true -->
```python
high_low_fare = None 

### BEGIN SOLUTION
high_low_fare = titanic_df.groupby(by='sex').agg({'fare':['max', 'min']})
### END SOLUTION
```

<!-- grade_id: cell-5d68da022e6dfa03, points: 1, grade: true, solution: false -->
```python
t.test_3j(high_low_fare)
```

**Question 3(k)**  <br> {points: 1}  

Use chaining to obtain the name of the youngest known female passenger and save it as object `youngest_female`.

<!-- grade_id: cell-2af735c6c0c8e9ed, points: 0, grade: false, solution: true -->
```python
youngest_female = None

### BEGIN SOLUTION
youngest_female = (titanic_df.groupby(by='sex')
                             .get_group('female')
                             .sort_values(by='age')
                             .iloc[0, 2]
                  )

### END SOLUTION

```

<!-- grade_id: cell-b1c0624507f4f69e, points: 1, grade: true, solution: false -->
```python
t.test_3k(youngest_female)
```

## 4. Titanic Analysis  


In exercise 3, we explored the dataset somewhat arbitrarily, trying to gain some insights into the data. In this exercise, we will carry out a number of steps to answer a single analysis question. 

> "How does a passenger's class, sex and age affect their chance of survival?"

Use chaining to answer each question.

One thing you may notice in the data is that there are missing values in the survived column that present themselves as`NaN`. This can be a potential problem as it could cause the analyst to draw inaccurate conclusions about the data. For this question, we will disregard this issue to concentrate on the skills we have learned up until now. We will, however, be returning to this topic further into the course and discuss how missing data effects analysis and different methods of handling them.

**Question 4(a)**  <br> {points: 1}  

Beginning the analysis with passenger class, find the mean survival rate of the 3 classes and plot this as a bar chart. Name the object `group_class`. Make sure that you give the plot a descriptive title and y-axis label.
You will need to groupby `pclass` and inorder to find the `mean` survival rate. Don't forget to `reset_index()` when plotting.

<!-- grade_id: cell-0390e77765430881, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# generate the plotting data
group_class_df = pd.DataFrame(titanic_df.groupby('pclass')['survived'].mean())

# plot the data
group_class = alt.Chart(
    group_class_df.reset_index(), width=500, height=300).mark_bar().encode(
        y=alt.Y('survived:Q', title="Survival Rate"),
        x=alt.X('pclass:N'),
    ).properties(title="Passenger Class Survival Rate")
### END SOLUTION
```

<!-- grade_id: cell-c73402b6c72374a4, points: 1, grade: true, solution: false -->
```python
t.test_4a(group_class)
```

Output:
```
'success'
```

**Question 4(b)**  <br> {points: 1}  

To answer the question, assign the letter associated with the correct answer to a variable in the code cell below:   
Which statement below is accurate?

A. Passengers from 1st class had the highest survival rate

B. Passengers from 2nd class had the highest survival rate 

C. Passengers from 3rd class had the highest survival rate


*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between `""`, assign the correct answer to an object called `answer4_b`*

<!-- grade_id: cell-76b9ead78ecfb0e0, points: 0, grade: false, solution: true -->
```python
answer4_b = None

### BEGIN SOLUTION
answer4_b = "A"
### END SOLUTION
```

<!-- grade_id: cell-7dbcb42fb864375a, points: 1, grade: true, solution: false -->
```python
t.test_4b(answer4_b)
```

**Question 4(c)**  <br> {points: 1}  

Plot the mean survival rates for each sex using a bar chart. Name the object `group_sex`. Make sure that you give the plot a descriptive title and y-axis label. You will need to groupby `pclass` and inorder to find the `mean` survival rate. Don't forget to `reset_index()` when plotting.    

<!-- grade_id: cell-3ec8f18303bb6b2f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# generate the plotting data
group_sex_df = pd.DataFrame(titanic_df.groupby('sex')['survived'].mean())

# plot the data
group_sex = alt.Chart(
    group_sex_df.reset_index(), width=500, height=300).mark_bar().encode(
        y=alt.Y('survived:Q', title="Survival Rate"),
        x=alt.X('sex'),
    ).properties(title="male and female passenger survival rate")

### END SOLUTION
```

<!-- grade_id: cell-f13869dd0bead0a2, points: 1, grade: true, solution: false -->
```python
t.test_4c(group_sex)
```

Output:
```
'success'
```

**Question 4(d)**  <br> {points: 1}  

To answer the question, assign the letter associated with the correct answer to a variable in the code cell below:   
Which statement below is accurate?

A. Male passengers had the highest survival rate

B. Female passengers had the highest survival rate 


*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between `""`, assign the correct answer to an object called `answer4_d`*

<!-- grade_id: cell-9149cc7bdba9811f, points: 0, grade: false, solution: true -->
```python
answer4_d = None

### BEGIN SOLUTION
answer4_d = "A"
### END SOLUTION
```

<!-- grade_id: cell-1b1831223a4df316, points: 1, grade: true, solution: false -->
```python
t.test_4d(answer4_d)
```

_Note to mike: Do you have a suggested better way of doing Question 4(e)?_ 

**Question 4(e)**  <br> {points: 1}  


Before we can do anything let's categorize the ages into 4 age classes into a column named `age_group`. Use the following conditions for each age class
- younger than 20 column values → "youth"   
- 20-39 column values → "adult"
- 40-59 column values → "older adult"   
- Age 60 and older column values → "senior" 

Find the mean survival rate of each `age_group` and plot this as a bar chart. Name the object `age_plot`.

<!-- grade_id: cell-440f649f81785f5c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
# generate plotting data
titanic_df.loc[titanic_df['age'] < 20, 'age_group'] = "youth"
titanic_df.loc[(titanic_df['age'] < 40) & (titanic_df['age'] >= 20),
               'age_group'] = "adult"
titanic_df.loc[(titanic_df['age'] < 60) & (titanic_df['age'] >= 40),
               'age_group'] = "older adult "
titanic_df.loc[titanic_df['age'] >= 60, 'age_group'] = "senior"
age_plot_df = titanic_df.groupby('age_group')['survived'].mean()

# plot data
age_plot = alt.Chart(
    age_plot_df.reset_index(), width=500, height=300).mark_bar().encode(
        y=alt.Y('survived:Q', title="Survival Rate"),
        x=alt.X('age_group'),
    ).properties(title="Age group survival rate")
### END SOLUTION

```

<!-- grade_id: cell-829cd55e54c6514b, points: 1, grade: true, solution: false -->
```python
t.test_4e(age_plot)
```

Output:
```
'success'
```

**Question 4(f)**  <br> {points: 1}  

To answer the question, assign the letter associated with the correct answer to a variable in the code cell below:   
Which statement below is accurate?

A. youth passengers had the highest survival rate

B. adult passengers had the highest survival rate 

C. older adult passengers had the highest survival rate 

D. senior passengers had the highest survival rate 


*Answer in the cell below using the uppercase letter associated with your answer. Place your answer between `""`, assign the correct answer to an object called `answer4_f`*


<!-- grade_id: cell-490bc519ab3b379d, points: 0, grade: false, solution: true -->
```python
answer4_f = None

### BEGIN SOLUTION
answer4_f = "A"
### END SOLUTION
```

<!-- grade_id: cell-c6b572c05ef2cec7, points: 1, grade: true, solution: false -->
```python
assert 'answer4_f' in locals(), "Your answer should be named 'answer4_f'"

# Note that this test has been hidden intentionally.
# It will provide no feedback as to the correctness of your answers.
# Thus, it is up to you to decide if your answer is sufficiently correct.
### BEGIN HIDDEN TESTS
def hidden_3(answer):
    assert not answer is None, "Your answer does not exist. Have you passed in the correct variable?"
    assert sha1((answer.lower() + "W").encode('utf8')).hexdigest(
    ) == "1953f1dae96593912b29962acc336ed1f8bb6600", "Answer is incorrect"
    return ("Success")

hidden_3(answer4_f)
### END HIDDEN TESTS
```

Output:
```
'Success'
```

**Question 4(g)**  <br> {points: 1}  

Let's go one step further now and do something a little new. Make a groupby object on 2 columns. Use columns `pclass` and `sex`. This is will give 6 separate groups: 
- Female and 1st class
- Female and 2nd class
- Female and 3rd class
- Male and 1st class
- Male and 2nd class
- Male and 3rd class

Next, find the mean value of the `survived` rate using `.loc[]` and sort them in decreasing survival rate. 

<!-- grade_id: cell-839b6becf2c070f7, points: 0, grade: false, solution: true -->
```python
group_class_age= None

### BEGIN SOLUTION
group_class_age = (titanic_df.groupby(
    ['pclass', 'sex']).mean().loc[:, 'survived'].sort_values(ascending=False))
### END SOLUTION
```

<!-- grade_id: cell-e72552bb0a719e28, points: 1, grade: true, solution: false -->
```python
t.test_4g(group_class_age)
```

**Question 4(h)**  <br> {points: 1}  
Now that you know how to group on multiple columns, We are going to try it on all 3 of the columns; `pclass`, `sex` and `age` . 

What is the maximum number of groups possible if we group by all 3 columns? 





*Answer in the cell below and assign the correct answer to an object called `answer4_h`*


<!-- grade_id: cell-f75154cca3c9749a, points: 0, grade: false, solution: true -->
```python
answer4_h = None

### BEGIN SOLUTION
answer4_h = 3 * 2 * 4
### END SOLUTION

```

<!-- grade_id: cell-9ccb283f50c544eb, points: 1, grade: true, solution: false -->
```python
t.test_4h(answer4_h)
```

**Question 4(i)**  <br> {points: 1}  


Groupby all three columns and find the mean value of the `survived` rate using `.loc[]`. Sort them in decreasing survival rate,  similarly to what you did in question 4(g)

<!-- grade_id: cell-f2f336f9cc068465, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
titanic_df
group_class_sex_age = (titanic_df.groupby(['pclass', 'sex', 'age'])
                             .mean()
                             .loc[:,'survived']
                             .sort_values(ascending=False)
                  )
### END SOLUTION
```

<!-- grade_id: cell-f60510077774b6c3, points: 1, grade: true, solution: false -->
```python
t.test_4i(group_class_sex_age)
```

**Question 4(j)**  <br> {points: 1}  

Which of the 24 groups had a 100% rate of survival? Answer by creating an object with the class, then sex, then age group of the highest survival rate group. For example, if it were 2nd class adult males, your answer would be `[2, 'male', 'adult']`


<!-- grade_id: cell-87140b36d2ef7f96, points: 0, grade: false, solution: true -->
```python
answer4_j = [2, 'male', 'adult']

### BEGIN SOLUTION
answer4_j = [1, 'female', 'senior']
### END SOLUTION
```

<!-- grade_id: cell-01fdec6b2e784aa3, points: 1, grade: true, solution: false -->
```python
t.test_4j(answer4_j)
```

Before submitting your assignment please do the following:

- Read through your solutions
- **Restart your kernel and clear output and rerun your cells from top to bottom** 
- Makes sure that none of your code is broken 
- Verify that the tests from the questions you answered have obtained the output "Success"

This is a simple way to make sure that you are submitting all the variables needed to mark the assignment. This method should help avoid losing marks due to changes in your environment.  

## Attributions
- City of Vancouver Open Data portal - [Street Trees Dataset](https://opendata.vancouver.ca/explore/dataset/street-trees/information/?disjunctive.species_name&disjunctive.common_name&disjunctive.height_range_id)
- MDS DSCI 523 - Data Wrangling course - [MDS's GitHub website](hhttps://ubc-mds.github.io/) 


## Module Debriefing

<video width="854" height="480" controls src="static/sample_video.mp4" />